Differential Regulation of Hepatic Macrophage Fate by Chi3l1 in MASLD

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

The manuscript by Shan et al seeks to define the role of the CHI3L1 protein in macrophages during the progression of MASH. The authors argue that the Chil1 gene is expressed highly in hepatic macrophages. Subsequently, they use Chil1 flx mice crossed to Clec4F-Cre or LysM-Cre to assess the role of this factor in the progression of MASH using a high fat high, fructose diet (HFFC). They found that loss of Chil1 in KCs (Clec4F Cre) leads to enhanced KC death and worsened hepatic steatosis. Using scRNA seq they also provide evidence that loss of this factor promotes gene programs related to cell death. From a mechanistic perspective they provide evidence that CHI3L serves as a glucose sink and thus loss of this molecule enhances macrophage glucose uptake and susceptibility to cell death. Using a bone marrow macrophage system and KCs they demonstrate that cell death induced by palmitic acid is attenuated by the addition of rCHI3L1. While the article is well written and potentially highlights a new mechanism of macrophage dysfunction in MASH and the authors have addressed some of my concerns there are some concerns about the current data that continue to limit my enthusiasm for the study. Please see my specific comments below.

Major:

(1) The authors' interpretation of the results from the KC (Clec4F) and MdM KO (LysMCre) experiments is flawed. The authors have added new data that suggests LyM-Cre only leads to a 40% reduction of Chil1 in KCs and that this explains the difference in the phenotype compared to the Clec4F-Cre. However, this claim would be made stronger using flow sorted TIM4hi KCs as the plating method can lead to heterogenous populations and thus an underestimation of knockdown by qPCR. Moreover, in the supplemental data the authors show that Clec4f-Cre x Chil1flx leads to a significant knockdown of this gene in BMDMs. As BMDMs do not express Clec4f this data calls into question the rigor of the data. I am still concerned that the phenotype differences between Clec4f-cre and LyxM-cre is not related to the degree of knockdown in KCs but rather some other aspect of the model (microbiota etc). It woudl be more convincing if the authors could show the CHI3L reduction via IF in the tissue of these mice.

We thank the reviewer for these constructive comments. We have performed FACSsorting of KCs (CD45+ F4/80hi CD11blow TIM4hi) or MoMFs (CD45+ F4/80low CD11bhi Ly6G- TIM4-) from Chil1fl/fl and Lyz2∆Chil1 or Clec4f∆Chil1mice, respectively. Compared with Chil1fl/fl mice, mRNA levels of Chil1 was reduced more than 90% in KCs from Clec4f∆Chil1 mice while not different in MoMFs (Revised Figure S3B). Besides, compared with Chil1fl/fl mice, mRNA levels of Chil1 was reduced more than 90% in MoMFs from Lyz2∆Chil1 mice while roughly 40% in KCs (Revised Figure S5B). This revised data support the phenotypic difference between Lyz2-CKO and Clec4f-CKO mice.

We agree with the reviewer that the significant knockdown of Chil1 in BMDM from Clec4f∆Chil1mice is confusing. To keep the rigor of our data, we remove this part from our manuscript.

Additionally, we performed immunofluorescence staining to detect Chi3l1 expression in liver tissues of these mice. The results show a reduction of Chi3l1 expression in KCs (TIM4+F4/80+ cells) of both Lyz2∆Chil1and Clec4f∆Chil1mice, with a more pronounced decrease in Clec4f∆Chil1mice (Author response image 1).

The expression of Chi3l1 in liver tissues of Chil1fl/fl, Lyz2∆Chil1and Clec4f∆Chil1mice.

Immunofluorescent staining to detect Chi3l1(green) expression in liver sections of Chil1fl/fl, Lyz2∆Chil1and Clec4f∆Chil1mice under normal chow diet. TIM4 (KCs marker, white), F4/80 (macrophage marker, red), nuclei were counterstained with DAPI, Scale bar=20 µm and 10 µm (Inset).

(2) Figure 4 suggests that KC death is increased with KO of Chil1. The authors have added new data with TIM4 tht better characterizes this phenotype. The lack of TIM4 low, F4/80 hi cells further supports that their diet model is not producing any signs of the inflammatory changes that occur with MASLD and MASH. This is also supported by no meaningful changes in the CD11b hi, F4/80 int cells that are predominantly monocytes and early Mdms). It is also concerning that loss of KCs does not lead to an increase in Mo-KCs as has been demonstrated in several studies (PMID37639126, PMID:33997821). This would suggest that the degree of resident KC loss is trivial.

We appreciate the reviewer’s insightful comment. We agree that our data show no substantial generation of monocyte-derived Kupffer cells (MoKCs) within the 16-week HFHC model. However, we do not believe the degree of resident KC loss is trivial, since 60% of KCs die at 16 weeks compared with 0 week (Revised Figure 5D). Instead, our observations align with a phased replacement model: recruited monocytes first differentiate into monocyte-derived macrophages (MoMFs), which we see accumulate (Revised Figure 5D), and only later adopt a KC phenotype. Consistent with this, our 16-week model shows significant EmKC loss and MoMFs expansion, but not yet the emergence of TIM4-MoKCs. This timing is supported by prior studies, where TIM4KCs were observed at 24 weeks, but not at 16 weeks, on similar diets (PMID: 33440159; PMID: 32888418). Therefore, we interpret our findings as capturing an earlier phase of MASLD progression, characterized by EmKC death and MoMF accumulation, prior to their full differentiation into MoKCs.

(3) The authors demonstrated that Clec4f-Cre itself was not responsible for the observed phenotype, which mitigates my concerns about this influencing their model.

We thank the reviewer for this comment and are pleased they agree that our control experiment using Clec4f-Cre alone confirms that the phenotype is specific to our genetic manipulation and not an artifact of the Cre driver.

(4) I remain somewhat concerned about the conclusion that Chil1 is highly expressed in liver macrophages. The author agrees that mRNA levels of this gene are hard to see in the datasets; however, they argue that IF demonstrates clear evidence of the protein, CHI3L. The IF in the paper only shows a high power view of one KC. I would like to see what percentage of KCs express CHI3L and how this changes with HFHC diet. In addition, showing the knockout IF would further validate the IF staining patterns.

We thank the reviewer for their thoughtful and constructive feedback. We agree that our initial conclusion regarding Chil1 expression in liver macrophages relied heavily on prior observations and was not sufficiently supported by the data presented. In response, we have revised our conclusion to state: "Hepatic macrophages express Chi3l1 and upregulate its expression following HFHC feeding." (Revised manuscript, page 4, line 136-137)

To strengthen this finding, we have replaced the original high-power image of a single Kupffer cell with a representative low-power view showing multiple F4/80+ macrophages (Revised Figure 1A). Furthermore, we performed quantitative colocalization analysis, which revealed that under normal chow diet (NCD), approximately 8% of F4/80+ macrophages are Chi3l1-positive. This proportion significantly increases to 15% upon HFHC feeding (Revised Figure 1A).

Additionally, to validate the specificity of the Chi3l1 immunofluorescence signal, we have included staining of liver sections from Chil1 knockout mice. In contrast to wildtype mice, Chi3l1 signal was completely absent within F4/80+ macrophages in Chil1-/- mice, confirming the specificity of the staining (Revised Figure 1B, Revised manuscript, page 4, line 152-157).

Minor:

(1) The authors have answered my question about liver fibrosis. In line with their macrophage data their diet model does not appear to induce even mild MASH.

We thank the reviewer for this observation. We agree that under our HFHC dietary conditions, the mice do not develop MASH pathology. However, we believe this earlystage model is a strength of our study, as it allows us to dissect the initial role of the Chi3l1-glucose interaction in regulating Kupffer cell fate during early MASLD, prior to the onset of significant fibrosis. This approach enables us to capture early macrophage adaptations (such as Chi3l1 upregulation) that might otherwise be masked or become secondary to the overt inflammation and scarring characteristic of late-stage MASH models.

Reviewer #2 (Public review):

In the revised version of the manuscript, the authors have attempted to address my questions, however, a number of my original concerns still remain.

Firstly, I had asked for a validation of the different CRE lines used - Lysm and Clec4f. The authors have now looked at BMDMs and KCs (steady state) from these animals. They conclude LysM only targets BMDMs not KCs, while CLEC4F targets both KCs and BMDMs. This I do not understand, BMDMs do not express CLEC4F so why are they targeted with this CRE? Additionally, BMDMs are not the correct control here, rather the authors should look at the incoming moMFs in the livers of these mice in the MASLD setting. Similarly, the KO in the MASLD KCs should be verified.

We thank the reviewer for these constructive comments. We have performed FACSsorting of KCs (CD45+ F4/80hi CD11blow TIM4hi) or MoMFs (CD45+ F4/80low CD11bhi Ly6G- TIM4-) from Chil1fl/fl and Lyz2∆Chil1 or Clec4f∆Chil1mice fed NCD or HFHC for 4 weeks, respectively. Compared with Chil1fl/fl mice, mRNA levels of Chil1 was reduced more than 90% in KCs from Clec4f∆Chil1 mice while not different in MoMFs at both 0 and 4 weeks (Revised Figure S3B). Besides, compared with Chil1fl/fI</sup mice, mRNA levels of Chil1fl/fI</sup was reduced more than 90% in MoMFs from Lyz2∆Chil1 mice while roughly 40% in KCs at both 0 and 4 weeks (Revised Figure S5B). This revised data support the phenotypic difference between Lyz2-CKO and Clec4f-CKO mice.

Then I had asked for validation of macrophage expression of Chil1 in other MASLD human and mouse datasets. The authors have looked into this, but the data provided do not suggest it is highly expressed by these cells either in the other mouse models or in the human. Nevertheless, they include a statement suggesting a similar expression pattern (although also being expressed by other cells). This is not an accurate discussion of the data and hence must be revised. This also prompted me to take another look at their data and this has left me querying the data in Figure 1D. Is the percent expressed 1%? In Figure 1C the scale goes from 0-100 but here 0-1. If we are talking about expression in 1% of cells which would fit with the additional public mouse data now analysed then how relevant are any of these claims? How sure are the authors that the effects seen are through KCs/moMFs? In figure 1D all cells profiled by scRNA-seq should be shown not just MFs to get a better sense of this data. What is macrophage expression of Chil1 compared with all other liver cells?

We thank the reviewer for the thoughtful feedback. We agree that the expression pattern of Chil1 should be described more accurately. To address this point, we examined four additional publicly available scRNA-seq datasets, including two mouse MASLD models and two human MASLD datasets (Author response image 2). Across these studies, the cell type with the highest Chil1 expression varied, whereas Chil1 transcripts were detected at relatively low frequency in macrophages (~1% of cells; Author response image 2C, E, K). To better present these data, we regenerated the UMAP plots to include all captured liver non-parenchymal cells, defined using the top two lineage specific markers (Author response image 3A–B). Consistent with Figure 2A–C, violin plots show that Chil1 is highly expressed in neutrophils, with only modest expression detected in macrophages (Author response image 3C). Further analysis of monocyte/macrophage subsets indicates that approximately ~1% of MoMFs or KCs express Chil1 (Author response image 3D–F). As the reviewer noted, the y-axis in Author response image 3F ranges from 0–1%, reflecting the low transcriptional detection frequency of Chil1 in macrophages, which is consistent with the additional public datasets analyzed.

We also recognize that mRNA detection by scRNA-seq does not necessarily reflect protein abundance. Therefore, we assessed Chi3l1 protein expression in hepatic macrophages using immunofluorescence staining for F4/80, TIM4, and Chi3l1 in liver sections from mice fed either normal chow diet (NCD) or HFHC diet. These analyses show that Chi3l1 protein is detectable in both KCs (TIM4+F4/80+) and MoMFs (TIM4-F4/80+) (Revised Figure 1A). Quantitative colocalization analysis revealed that under NCD conditions, approximately 8% of F4/80+ macrophages are Chi3l1-positive, which increases to ~15% following HFHC feeding (Revised Figure 1A). To confirm antibody specificity, we additionally performed staining in Chil1 knockout mice. In contrast to wild-type mice, Chi3l1 signal was completely absent in F4/80+ macrophages from Chil1-/- mice, validating the specificity of the staining (Revised Figure 1B). Together, these results suggest that low-abundance Chil1 transcripts may be under-detected by scRNA-seq, whereas immunofluorescence captures accumulated protein. Importantly, our functional experiments using Clec4f-Cre– mediated deletion directly support that the observed phenotypes are mediated through Kupffer cells, regardless of expression levels in other liver cell types.

In response to the reviewer’s comments, we have made the following revisions:

(1) Softened our conclusion to: “Hepatic macrophages express CHI3L1 and upregulate its expression following HFHC feeding” (Revised manuscript, page 4, lines 136–137).

(2) Included representative low-magnification images showing multiple F4/80+ macrophages along with quantitative analysis (Revised Figure 1A).

(3) Added immunofluorescence staining of Chil1-/- liver sections demonstrating complete absence of Chi3l1 signal in F4/80+ macrophages, validating antibody specificity (Revised Figure 1B).

(4) Regenerated UMAP plots to display all liver non-parenchymal cells and clearly indicate the low detection frequency of Chil1 transcripts in macrophages (Author response image 3).

(5) Revised the relevant text to more accurately describe Chil1 expression patterns in hepatic macrophages (Revised manuscript, page 4, lines 136–157).

Analysis of Chil1 expression in additional single-cell RNA sequencing datasets.

(A-C) Chil1 expression in a mouse model of NASH. (A) t-SNE projection of cell clusters from scRNA-seq data (GSE1283338) of livers from C57BL/6J mice fed a control or NASH diet for 30 weeks. (B) Dot plot showing scaled Chil1 expression across all identified cell clusters. (C) Dot plot of scaled Chil1 expression after excluding the neutrophil cluster, highlighting expression in macrophage populations. Analyzed cell clusters and cell numbers: KC_H (healthy, 1178); KC3_Control (1142); KC_N (NASH, 1045); KN_RM (recruited macrophage in KC niche, 950); Proliferating_KC (364); PDC_Control (356); Ly6CHi_RM (320); LSEC (299); NK_NKT (393); B_cell (244); DC_1 (107); DC_2 (118); Ly6CLo_RM (127); Hepatocyte (57); PDC_NASH (46); Neutrophil (21). (D-E) Chil1 expression during NAFLD progression in a mouse Western diet model. (D) t-SNE projection of cell clusters from scRNA-seq data (GSE156059) of livers from C57BL/6J mice fed a Western diet with fructose/sucrose for 12, 24, and 36 weeks. (E) Dot plot showing scaled Chil1 expression across all identified cell clusters. Analyzed cell clusters and cell numbers: capsule macs (250), LAMs (1419), Ly6chi monocytes (6912), mac1 (638), moKCs (767), Patrolling monocytes (690), Prolif.macs (521), Resident KCs (3629), Transitioning monocytes (3615). (F-H) Chil1 expression in human cirrhotic liver biopsies. (F) t-SNE projection of cell clusters from scRNA-seq data (GSE136103) of healthy and cirrhotic human liver samples. (G) Dot plot showing scaled Chil1 expression across major cell lineages. (H) Dot plot of scaled Chil1 expression specifically within the mononuclear phagocyte (MP) population. Analyzed cell clusters and cell numbers: B cell (1951); cycling (967); Epithelia (3751); ILC (10091); mast cell (2511); Mesenchyme (2382); MP (10874); pDC (317); Plasma cell (877); T cell (19076). (I-K) Chil1 expression in a human NAFLD explant. (I) t-SNE projection of cell clusters from scRNA-seq data (GSE190487) of a human NAFLD liver explant. (J) Dot plot showing scaled Chil1 expression across all identified cell clusters. (K) Dot plot of scaled Chil1 expression within the MP subpopulations. Analyzed cell clusters and cell numbers: B cell (1278); Cycling (152); MP (2897); pDC (391); Plasma cell (85); T cell (1551); KC (403); SAMac (scar-associated macrophages, 723); TM (tissue monocytes, 1265).

Hepatic macrophages express Chi3l1.

(A-D) Wildtype C57BL/6J mice were fed either a normal chow diet (NCD) or HFHC for 16 weeks. NPCs were isolated and subjected to BD Rhapsody scRNA sequencing. (A) Uniform manifold approximation and projection (UMAP) plots illustrate the clustering of NPCs from the livers of mice fed NCD and HFHC. Major cell types are colored. (B) Heatmap showing the mean expression of top2 markers of each cell type. (C) Violin plots show the RNA expression of Chil1 between NCD and HFHC livers in each cell cluster. (D) UMAP plots depict the clustering of Monocytes/Macrophages in the livers of mice fed NCD and HFHC. Cell clusters are color-coded. (E) Dot plot displays the scaled gene expression levels of lineage-specific marker genes in different cell clusters. (F) Dot plot shows the scaled gene expression levels of Chil1 in the indicated cell clusters.

The cell death had also previously concerned me that 40-60% of KCs were tunel +ve. I do not understand how 60% are +ve at 8 weeks but then they have more or less same number of TIM4+ cells at 16 weeks? How can this be? why do the tunel +ve cells not die? This concern remains as I don't understand how they reached these numbers given the images. Additional, larger images were also not provided to be sure that they are representative images in the figure. Now in the images provided, there are clearly cells which are TIM4+ where the tunel does not overlap, likely it is in a LSEC or other neighbouring cell. Indeed also taking Fig S11b as an example there are ˜7KCs and at best 1 expresses tunel so how do they get to 60%?

We thank the reviewer for these constructive feedback. We agree that the sustained TUNEL positivity without corresponding KC depletion presents an apparent paradox. Based on our data, we propose that TUNEL-positive KCs represent cells in a prolonged stressed or pre-apoptotic state rather than undergoing immediate clearance. This interpretation is supported by the relatively stable TIM4+ cell numbers between 8 and 16 weeks, which would be inconsistent with rapid cell death and removal. Previous studies (PMID: 33440159; PMID: 32888418) have similarly documented gradual KC loss during MASLD progression, supporting our view that KC death occurs over an extended timeframe rather than acutely.

Regarding quantification concerns, we acknowledge that the representative images in the original figure may have been misleading. To address this, we have now quantified KC apoptosis using low-magnification fields across multiple liver sections to ensure statistical rigor. Figure S11B (now Revised Figure S9B) presents these data, showing that under NCD conditions, KC apoptosis rates are minimal in both genotypes. Following HFHC feeding, apoptosis rates are comparable between Chil1fl/fl and Lyz2Δ Chil1 mice. Importantly, we have replaced all TIM4/TUNEL co-staining images with lowmagnification representative images in the revised figures (Revised Figure 1A, 1B, 5E, S9A, S9B). These images better reflect the quantitative data and confirm that the originally highlighted high-magnification fields were not representative of global apoptosis rates.

Reviewer #3 (Public review):

This paper investigates the role of Chi3l1 in regulating the fate of liver macrophages in the context of metabolic dysfunction leading to the development of MASLD. I do see value in this work, but some issues exist that should be addressed as well as possible.

Here are my comments:

(1) Chi3l1 has been linked to macrophage functions in MASLD/MASH, acute liver injury, and fibrosis models before (e.g., PMID: 37166517), which limits the novelty of the current work. It has even been linked to macrophage cell death/survival (PMID:31250532) in the context of fibrosis, which is a main observation from the current study.

We thank the reviewer for raising this important point and acknowledge previous studies linking Chi3l1 to macrophage function in liver disease. However, several aspects of our work extend beyond these prior reports. First, although global Chi3l1 deficiency has been shown to promote macrophage apoptosis in toxin-induced fibrosis models (PMID: 31250532), our study demonstrates that Chi3l1 differentially regulates the fate of distinct hepatic macrophage subsets embryo-derived Kupffer cells (KCs) and monocyte-derived macrophages (MoMFs)—in MASLD. To our knowledge, this subset-specific regulation of hepatic macrophages has not been previously described. Second, we identify a previously unrecognized metabolic mechanism by which Chi3l1 regulates macrophage survival. Specifically, we find that Chi3l1 binds glucose and promotes glucose uptake, thereby protecting the highly glucose-dependent KCs from metabolic stress–induced death, while exerting minimal effects on MoMFs. This mechanism is distinct from the previously reported Fas/Akt-mediated pathway (PMID: 31250532) and highlights a metabolic checkpoint controlling macrophage subset– specific vulnerability. Third, our findings reveal context- and cell type-dependent roles of Chi3l1. While myeloid-specific deletion of Chi3l1 has been reported to ameliorate steatohepatitis and fibrosis (PMID: 37166517), our KC-specific deletion model shows that loss of Chi3l1 in KCs exacerbates disease, indicating a previously unrecognized protective role of Chi3l1 in KCs during early MASLD. Together, these findings provide new insights into macrophage subset-specific regulation, identify a novel glucose related metabolic mechanism, and reveal context-dependent functions of Chi3l1 in MASLD pathogenesis.

(2) The LysCre-experiments differ from experiments conducted by Ariel Feldstein's team (PMID: 37166517). What is the explanation for this difference? - The LysCre system is neither specific to macrophages (it also depletes in neutrophils, etc), nor is this system necessarily efficient in all myeloid cells (e.g., Kupffer cells vs other macrophages). The authors need to show the efficacy and specificity of the conditional KO regarding Chi3l1 in the different myeloid populations in the liver and the circulation.

We thank the reviewer for raising this important point regarding the specificity of the genetic models and the apparent discrepancy with the study by Feldstein and colleagues (PMID: 37166517). To address these concerns, we performed additional experiments to directly assess the efficiency and cell-type specificity of Chi3l1 deletion in our models.

(1) Efficiency and specificity of LysM-Cre and Clec4f-Cre models

We isolated KCs (CD45+ F4/80hi CD11blow TIM4hi) or MoMFs (CD45+ F4/80low CD11bhi Ly6G- TIM4-) by FACS from Chil1fl/fl, Lyz2∆Chil1 and Clec4f∆Chil1mice fed either NCD or HFHC diet. Consistent with the known specificity of these Cre lines, Clec4f-Cre resulted in >90% reduction of Chil1 mRNA in KCs with no significant change in MoMFs (Revised Figure S3B), confirming efficient KC-specific deletion. In contrast, LysM-Cre reduced Chil1 expression by >90% in MoMFs but only ~40% in KCs (Revised Figure S5B). These data support the reviewer’s concern that LysM-Cre mediates incomplete recombination in KCs, whereas the Clec4f-Cre model provides KC-specific deletion, explaining why the phenotype observed in Lyz2∆Chil1 mice is relatively modest.

(2) Relationship to the study by Feldstein et al.

We agree that our LysM-Cre results appear different from those reported by Feldstein and colleagues. However, considering the new recombination data and differences in disease models, we believe the findings are complementary rather than contradictory. First, the disease models differ substantially. Feldstein et al. used a CDAA-HFAT diet for 10 weeks, which rapidly induces severe inflammation and fibrosis, whereas our study employed a long-term HFHC diet, modeling the more gradual metabolic progression of MASLD. These distinct disease contexts may engage different CHI3L1dependent pathways. Second, the mechanistic focus differs. Feldstein et al. reported that myeloid Chi3l1 promotes steatohepatitis and fibrosis through inflammatory macrophage recruitment and IL13Rα2-mediated stellate cell activation. In contrast, our study identifies a metabolic mechanism in which CHI3L1 binds glucose and promotes glucose uptake, protecting metabolically vulnerable KCs from stress-induced death. Finally, and importantly, KC-specific deletion using Clec4f-Cre recapitulates the key phenotypes observed in our study, including effects on KC survival and metabolic regulation. This confirms that the observed effects are KC-autonomous and not due to broader Cre activity in other myeloid populations.

Together, these additional experiments clarify the recombination efficiency of our models and demonstrate that our conclusions are supported by KC-specific genetic evidence.

(3) The conclusions are exclusively based on one MASLD model. I recommend confirming the key findings in a second, ideally a more fibrotic, MASH model.

We thank the reviewer for this valuable suggestion. To address this point, we tested our key findings in an additional MASH model using a methionine–choline-deficient (MCD) diet. First, we examined Chi3l1 expression in this model. Wild-type mice fed an MCD diet for 6 weeks showed significantly increased Chi3l1 mRNA and protein levels in liver tissues compared with NCD controls, confirming diet-induced upregulation (Revised Figure 3A–B). To determine the functional contribution of Kupffer cell–derived Chi3l1, we subjected Clec4fΔChil1 mice and Chil1fl/fl controls to MCD feeding for 6 weeks. Body weight was comparable between genotypes throughout the feeding period (Revised Figure 3C). However, KC-specific deletion of Chi3l1 significantly exacerbated MCD diet–induced liver pathology, including increased steatosis, inflammation, and fibrosis, as indicated by higher MASLD activity scores, enhanced Oil Red O staining, increased Sirius Red deposition, and elevated α-SMA expression (Revised Figure 3D). Consistent with these histological findings, Clec4fΔChil1 mice exhibited an increased liver index, whereas serum ALT levels remained comparable between groups, suggesting increased hepatic lipid accumulation rather than aggravated hepatocellular injury (Revised Figure 3E). In addition, serum and hepatic triglyceride levels and serum cholesterol were significantly elevated, while hepatic cholesterol levels were not significantly different from controls (Revised Figure 3E). Together, these results validate our findings in an independent MASH model and further support a protective role for Kupffer cell–derived Chi3l1 in limiting steatosis and disease progression (Revised manuscript, page 5, line 188-205).

(4) Very few human data are being provided (e.g., no work with own human liver samples, work with primary human cells). Thus, the translational relevance of the observations remains unclear.

We thank the reviewer for raising this important point. We agree that additional human validation would further strengthen the translational relevance of our findings. We initially attempted to examine macrophage cell death in human liver samples by performing TUNEL and F4/80 co-staining on human liver cancer tissues. However, we did not detect clear colocalization in these samples. We speculate that this may reflect differences in disease context and stage, as the available samples represent endstage liver disease, whereas our study focuses on early MASLD progression. Despite this limitation, we provide several lines of evidence supporting the human relevance of our findings. First, analysis of multiple public human MASLD scRNA-seq datasets demonstrates Chi3l1 expression in hepatic macrophages (Figure 2F–K). Second, analysis of public bulk RNA-seq datasets shows that Chi3l1 expression positively correlates with MASLD disease activity and progression (Revised Figure 1EF). Third, our observations are consistent with previous clinical studies reporting elevated CHI3L1 levels in patients with MASLD/MASH and advanced liver disease. We acknowledge that functional validation in primary human macrophages or human liver tissues would further strengthen the translational significance of this work. This limitation and future direction have now been added to the Discussion (Revised manuscript, page 10, lines 409–411).

Comments on revisions:

The authors have done a thorough job addressing my comments. However, I am not convinced about the MCD diet model, which is somewhat hidden in the Supplementary Files. Neither seems MASH different nor are any fibrosis data shown to support the conclusions. I am not satisfied with this part of the revised manuscript, and I do not agree that the second MASH model would support the conclusions.

We thank the reviewer for their continued careful evaluation and for highlighting the need for clearer presentation of the MCD model data. To address this concern, we have substantially revised this section of the manuscript. First, the MCD model results have now been moved from the Supplementary Figure to a new main figure (Revised Figure 3) to improve visibility and clarity. Second, we have added additional fibrosis analyses, including Sirius Red staining and α-SMA immunostaining, to directly assess fibrotic changes. These analyses show that MCD feeding induces significant collagen deposition in control mice and that fibrosis is further increased in Clec4fΔChil1 mice (Revised Figure 3D). Importantly, the MCD model recapitulates the key phenotypes observed in the HFHC model, with KC-specific Chi3l1 deletion leading to increased MASLD progression. These findings support the conclusion that the protective role of Kupffer cell–derived Chi3l1 is not restricted to a single dietary model, but is observed across distinct models of steatohepatitis. We hope that these revisions clarify the results and strengthen the evidence supporting our conclusions.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

Minor:

Line 73 - should be moMfs not moKCs

We thank the reviewer for this helpful comment. The term moKCs was used intentionally in line 73 to refer to monocyte-derived Kupffer cells, rather than MoMFs (monocyte-derived macrophages). To avoid potential confusion, we have clarified the terminology in the revised manuscript.

Methods: diet is mentioned for 6 weeks but for HFHC should be 16.

The correction has been made in the Methods section (page 3，line115).

Liver/body weight ratios are >3 then I think it is body/liver weight ratio?

We thank the reviewer for this query. The reported values represent liver-to-body weight ratios, calculated as (liver weight ÷ body weight) × 100%. A value of ~3% is consistent with the expected range for mice with MASLD-associated hepatomegaly.

This clarification has been added to the revised figure legend.

Figure 5F - what happens in Clec4f-CRE mice fed HFHC?

We thank the reviewer for this question. Western blot analysis showed that the HFHC diet upregulated Chi3l1 protein in the livers of Clec4f-Cre mice post HFHC diet (Author response image 4.), similar to the increase observed in wild-type mice.

The expression of Chi3l1 in serum of Clec4f cre mice.

(A) Western blot to detect Chi3l1 expression in murine serum of Clec4f cre mice before and after HFHC feeding. n=3 mice/group.

Changes from the earlier version to the revised version

Abstract

Metabolic dysfunction-associated steatotic liver disease (MASLD) progression involves the replacement of protective embryo-derived Kupffer cells (KCs) by inflammatory monocyte-derived macrophages (MoMFs), yet the regulatory mechanisms remain unclear. Here, we identify chitinase 3-like 1 (Chi3l1/YKL-40) as a critical metabolic regulator of hepatic macrophage fate. We observed high expression of Chi3l1 in both KCs and MoMFs during MASLD development. Genetic deletion of Chi3l1 specifically in KCs significantly exacerbated MASLD severity and metabolic dysfunction, whereas MoMF-specific Chi3l1 deletion showed minimal metabolic effects. Mechanistic studies revealed that this cell type-specific regulation arises from differential metabolic requirements: KCs display elevated glucose metabolism compared to MoMFs. Chi3l1 directly interacts with glucose to inhibit its cellular uptake, thereby selectively protecting glucose-dependent KCs from metabolic stress-induced cell death while having negligible effects on less glucose-dependent MoMFs. These findings uncover a novel Chi3l1-mediated metabolic checkpoint that preferentially maintains KCs populations through glucose metabolism modulation, providing important new insights into the pathogenesis of MASLD and potential therapeutic strategies targeting macrophage-specific metabolic pathways.

Introduction

Earlier version: Metabolic dysfunction-associated steatotic liver disease (MASLD) has become the most prevalent chronic liver disorder in western populations, affecting approximately 30% of adults and driven by its strong association with obesity and metabolic syndrome1. The disease spectrum ranges from matabolic dysfunction-associated fatty liver (MAFL) to metabolic dysfunction-associated steatohepatitis (MASH), with the latter characterized by steatosis, inflammation, hepatocyte ballooning, and progressive fibrosis2. Central to MASLD pathogenesis are hepatic macrophages, particularly the embryo-derived Kupffer cells (KCs) that reside in liver sinusoids3,4. These self-renewing resident macrophages5, play crucial roles in lipid homeostasis, as evidenced by studies showing that depletion of CD207+ KCs leads to impaired triglyceride storage6. As MASH progresses, dying KCs are progressively replaced by monocyte-derived macrophages (MoMFs) that exhibit heightened inflammatory properties and contribute to liver damage6,7. For example, one study demonstrated that in diet-induced MASH, KCs enhancer landscapes and gene expression profiles are profoundly reprogrammed (including up-regulation of Trem2 and Cd9) and KCs identity is lost, while MoMFs adopt convergent epigenomes, transcriptomes and functions during macrophage recruitment and adaptation in MASH8. Another work showed that in MASLD the number of resident KCs declines and MoMFs accumulate; these recruited macrophages include subsets that either mirror homeostatic KCs or resemble lipid-associated macrophages (LAMs) from obese adipose tissue, with the LAM-type expressing osteopontin and localizing to fibrotic zones9. Together, these findings highlight that this transition from protective embryo-derived KCs (EmKCs) to inflammatory monocyte-derived KCs (MoKCs) represents a critical juncture in disease progression, yet the mechanisms regulating this shift remain poorly understood.

Revised version: Metabolic dysfunction-associated steatotic liver disease (MASLD) has become the most prevalent chronic liver disorder in western populations, affecting approximately 30% of adults and driven by its strong association with obesity and metabolic syndrome1. The disease spectrum ranges from matabolic dysfunction-associated fatty liver (MAFL) to metabolic dysfunction-associated steatohepatitis (MASH), with the latter characterized by steatosis, inflammation, hepatocyte ballooning, and progressive fibrosis2. Central to MASLD pathogenesis are hepatic macrophages, particularly the embryo-derived Kupffer cells (KCs) that reside in liver sinusoids3,4. These self-renewing resident macrophages5, play crucial roles in lipid homeostasis, as evidenced by studies showing that depletion of CD207+ KCs leads to impaired triglyceride storage6. As MASH progresses, dying KCs are progressively replaced by monocyte-derived macrophages (MoMFs) that exhibit heightened inflammatory properties and contribute to liver damage6,7. For example, one study demonstrated that in diet-induced MASH, KCs enhancer landscapes and gene expression profiles are profoundly reprogrammed (including up-regulation of Trem2 and Cd9) and KCs identity is lost, while MoMFs adopt convergent epigenomes, transcriptomes and functions during macrophage recruitment and adaptation in MASH8. Another work showed that in MASLD the number of resident KCs declines and MoMFs accumulate; these recruited macrophages include subsets that either mirror homeostatic KCs or resemble lipid-associated macrophages (LAMs) from obese adipose tissue, with the LAM-type expressing osteopontin and localizing to fibrotic zones9. Together, these findings highlight that this transition from protective embryo-derived KCs (EmKCs) to monocyte-derived KCs (MoKCs) represents a critical juncture in disease progression, yet the mechanisms regulating this shift remain poorly understood.

A key determinant of macrophage function is cellular metabolism. Macrophages dynamically switch between glycolytic and oxidative phosphorylation pathways to adapt to environmental changes10. During MASLD, hepatic macrophages increase their glycolytic activity, which may exacerbate inflammation and tissue damage11–13. While glucose metabolism is known to influence macrophage polarization, its specific role in determining hepatic macrophage fate - particularly the balance between KCs and MoMFs - remains unknown. Chitinase 3-like 1 (Chi3l1/YKL-40) has emerged as an important regulator of macrophage biology, promoting cell survival through ERK1/2 and PI3K/Akt pathways while modulating anti-inflammatory cytokines like IL-1014–18. However, its potential role in macrophage metabolic reprogramming, particularly in the context of hepatic glucose metabolism, has not been explored.

In this study, we identify a novel mechanism by which Chi3l1 governs hepatic macrophage fate through metabolic regulation. We demonstrate that Chi3l1 directly interacts with glucose to suppress its uptake in macrophages. Strikingly, this interaction selectively protects glucose-high KCs from cell death in MASLD conditions, while having minimal effect on glucose-low MoMFs. These findings reveal a previously unrecognized Chi3l1-mediated metabolic checkpoint that maintains KC populations, providing new insights into the pathogenesis of MASLD and potential therapeutic strategies.

Materials and methods

Animal experiments and procedures

Added in the revised version: Animals

Earlier version: Animals Chil1-/- (strain no. T014402), Chil1flox//flox (strain no. T013652), Lyz2-cre (strain no. T003822), Clec4f-cre (strain no. T036801) with a C57BL/6J background were purchased from GemPharmatech. Rosa tdtomato mice (strain no. C001181) were purchased from Cyagen. Accordingly, C57BL/6J mice (strain no. N000013) were used as wild-type (WT) mice. To generate Clec4f△Chil1 mice, Chil1flox//flox mice were crossed with Clec4f-cre mice and knock out efficiency was examined in KCs or BMDM by qRT-PCR and western (Figure S4B, C). To generate Clec4fRosa tdtomato mice, Rosa tdtomato mice were crossed with Clec4f-cre mice to examine the expression specificity of Clec4f-cre (Figure S4D). To generate Lyz2△Chil1 mice, Chil1flox//flox mice were crossed with Lyz2-cre mice and knock out efficiency was examined in BMDM by western blot and qRT-PCR (Figure S7B, C). All mouse colonies were maintained at the Animal Core Facility of Yunnan University. The animal studies were approved by the Yunnan University Institutional Animal Care and Use Committee (IACUC, Approval No. YNU20220314). Male mice aged 6-8 weeks were used in this study.

Revised version: Chil1-/-(strain no. T014402), Chil1flox//flox (strain no. T013652), Lyz2-cre (strain no. T003822), Clec4f-cre (strain no. T036801) with a C57BL/6J background were purchased from GemPharmatech. Rosa tdtomato mice (strain no. C001181) were purchased from Cyagen. Accordingly, C57BL/6J mice (strain no. N000013) were used as wild-type (WT) mice. To generate Clec4f△Chil1 mice, Chil1flox//flox mice were crossed with Clec4f-cre mice. To generate Clec4fRosa tdtomato mice, Rosa tdtomato mice were crossed with Clec4f-cre mice. To generate Lyz2△Chil1 mice, Chil1flox//floxmice were crossed with Lyz2-cre mice. Male mice have been the choice in the vast majority of the studies of MASLD reported in the literature19,20. Therefore, we used male mice in the majority of the experiments presented. All mouse colonies were maintained at the Animal Core Facility of Yunnan University. The animal studies were approved by the Yunnan University Institutional Animal Care and Use Committee (IACUC, Approval No. YNU20220314).

Added in the revised version: Construction of MASLD/MASH mouse model

Earlier version: Construction of MASLD/MASH mouse model Mice were provided a high-fat and high-cholesterol diet (Research Diet, d12108c, 40 kcal% fat and 1.25% cholesterol) or a methionine and choline deficient diet (Research Diet, A02082002BR) for 6 weeks. Throughout the feeding period, the body weight and food consumption of the mice were observed and recorded weekly. Once the dietary intervention was completed, the mice were euthanized. Liver and murine serum samples were collected for further analysis. Alanine aminotransferase (ALT) and aspartate aminotransferase (AST) levels in the serum, as well as cholesterol (TC) and triglyceride (TG) levels in both serum and liver tissues, were quantified using commercially available kits (Nanjing Jiancheng Bioengineering Institute).

Revised version: Mice were provided a high-fat and high-cholesterol diet (Research Diet, d12108c, 40 kcal% fat and 1.25% cholesterol) for 16 weeks or a methionine and choline deficient diet (Research Diet, A02082002BR) for 6 weeks. Throughout the feeding period, the body weight and food consumption of the mice were observed and recorded weekly. Once the dietary intervention was completed, the mice were euthanized. Liver and murine serum samples were collected for further analysis. Alanine aminotransferase (ALT) and aspartate aminotransferase (AST) levels in the serum, as well as cholesterol (TC) and triglyceride (TG) levels in both serum and liver tissues, were quantified using commercially available kits (Nanjing Jiancheng Bioengineering Institute).

Statistical analysis

Data are presented as mean ± standard error of the mean (SEM) in all graph figures. Statistical analyses were conducted using the SPSS statistics software (Version 22). To compare the two groups, an unpaired two-tailed Student’s t-test was used. One-way analysis of variance (ANOVA) was performed for comparisons involving three or more groups. For patients with MASLD liver, the samples were tested using the Mann-Whitney test. Statistical significance was set at p < 0.05 and p value is indicated. All cell culture results represent at least three independent experiments.

Additional Methods

Additional detailed methods can be found in the Supporting Information.

Results

Added in the revised version: Hepatic macrophages express Chi3l1 and upregulate its expression post HFHC diet

Added in the revised version: In our previous study, we found that Chi3l1 expressed by hepatic macrophages influences macrophage function during acute liver injury21. Therefore, we sought to determine whether Chi3l1 also plays a role in MASLD and whether its expression in hepatic macrophages is altered in this context. To this end, we first established a mouse model of MASLD by feeding C57BL/6J wild-type mice a normal chow diet (NCD) or an high-fat high-cholesterol (HFHC) diet for 16 weeks. Histological analysis of liver sections using Hematoxylin and Eosin (H&E) and Sirius Red staining revealed marked lipid accumulation without apparent fibrosis (Figure S1A). Consistently, western blot analysis showed no upregulation of α-SMA (Figure S1B), suggesting that our HFHC model represents an early stage of MASLD.

Added in the revised version: Next, we performed immunofluorescence staining for Chi3l1 in liver sections using antibodies against TIM4 (a KCs marker), F4/80 (a pan-macrophage marker), and Chi3l1. This confirmed that Chi3l1 is expressed in both KCs (TIM4+F4/80+ cells) and MoMFs (TIM4-F4/80+ cells), with elevated expression under HFHC conditions (Figure 1A). To validate the specificity of Chi3l1 staining, we generated Chil1-/- mice and confirmed knockout efficiency by qRT-PCR (Figure S2A, B). Immunofluorescence staining for Chi3l1 in liver sections from WT and Chil1-/- mice showed that the anti-Chi3l1 antibody specifically detected Chi3l1 in WT but not Chil1-/- mice, confirming the specificity of the staining (Figure 1B). We next assessed whether Chi3l1 is upregulated by HFHC feeding by measuring its protein levels in isolated KCs and whole liver tissue via western blotting. A marked increase in Chi3l1 expression was observed in both KCs and liver tissue following HFHC diet feeding (Figure 1C, D). Consistently, patients with MAFL or MASH exhibited elevated hepatic Chi3l1 mRNA levels, which correlated with MASLD severity and fibrosis stage (Figure 1E, F). These findings suggest that Chi3l1 is expressed in hepatic macrophages and may contribute to MASLD progression.

Added in the revised version: Hepatic macrophages express Chi3l1 and upregulate its expression post HFHC diet.

Removed in the revised version: Hepatic macrophages express Chi3l1

Removed in the revised version: To investigate the dynamic changes in hepatic macrophages, we performed BD Rhapsody scRNA-seq on non-parenchymal cells (NPCs) isolated from healthy livers of mice fed a normal chow diet (NCD) and from livers of mice with MASLD induced by a high-fat, high-cholesterol (HFHC) diet. Prior to NPC isolation for scRNA-seq, liver sections were subjected to H&E and Sirius Red staining, which revealed marked lipid accumulation without apparent fibrosis (Figure S1A). Consistently, western blot analysis of α-SMA showed no upregulation (Figure S1B), confirming that our HFHC model represents an early stage of MASLD. After quality control and filtration, we retained 23,312 cells from NCD livers and 6,567 cells from HFHC livers for downstream analysis. Using a graph-based clustering approach, we identified 32 distinct cell populations, visualized via uniform manifold approximation and projection (UMAP) (Figure 1A). Monocyte/macrophage subsets were further defined based on lineage-specific markers: Monocytes expressed Ly6c2, Chil3, S100a6, Ccr2, Itgam, and Cx3cr1 but lacked macrophage markers. KCs were marked by Cd68, Vsig4, Clec4f, TIM4, Adgre1, and Clec1b. MoMFs were negative for KCs markers but positive for macrophage markers such as Ccr2, Cx3cr1, Cd9, Itgax, Gpnmb, Cd68, and Adgre1 (Figure 1B, C; UMAP in Fig S2A)8,19. ScRNA-seq analysis revealed high expression of Chil1 (Chil1 being the gene name for Chi3l1) in hepatic macrophages (Figure 1D). To validate this, we performed immunofluorescence staining for Chi3l1 in hepatic macrophages using antibodies against TIM4, F4/80, and Chi3l1. Our results confirmed Chi3l1 expression in hepatic macrophages under both NCD and HFHC conditions (Figure 1E). To determine whether Chi3l1 expression is upregulated following HFHC diet feeding, we analyzed Chi3l1 protein levels in isolated KCs and whole liver tissue by western blotting. The results revealed a marked increase in Chi3l1 expression in both KCs and liver after HFHC diet (Figure 1F, G). To strengthen our findings, we further examined four additional publicly available scRNA-seq datasets—two from mouse models and two from human MASLD patients (Figure S3). Across these datasets, the specific cell type showing the highest Chil1 expression varied somewhat between studies, likely reflecting model differences and disease stages. Nevertheless, Chil1 expression was consistently enriched in hepatic macrophage populations, including both Kupffer cells and infiltrating macrophages, in mouse and human livers. Notably, Chil1 expression was higher in MoMFs compared to resident KCs, supporting its upregulation during MASLD progression. Consistently, patients with MAFL or MASH exhibited elevated hepatic Chil1 mRNA levels, which correlated with MASLD severity and fibrosis stage (Figure 1H, I). These findings suggest that Chi3l1 is expressed in hepatic macrophages and may play a role in MASLD progression.

Removed in the revised version: Hepatic macrophages express Chi3l1.

Earlier version: (A-D) Wildtype C57BL/6J mice were fed either a normal chow diet (NCD) or HFHC for 16 weeks. NPCs were isolated and subjected to BD Rhapsody scRNA sequencing. (A) Uniform manifold approximation and projection (UMAP) plots illustrate the clustering of NPCs in the livers of mice fed NCD and HFHC. Cell clusters are color-coded, with monocytes/macrophages clusters outlined. (B) UMAP plots depict the clustering of Monocytes/Macrophages in the livers of mice fed NCD and HFHC. Cell clusters are color-coded. (C) Dot plot displays the scaled gene expression levels of lineage-specific marker genes in different cell clusters. (D) Dot plot shows the scaled gene expression levels of Chil1 in the indicated cell clusters. (E) Immunofluorescent staining of TIM4 (white), F4/80 (red), Chi3l1 (green), and nuclear DAPI (blue) in liver sections of mice fed with either NCD or HFHC for 16 weeks, illustrating Chi3l1 expression in hepatic macrophages. Scale bar=20μm and 5μm (zoom). (F, G) Western blot analysis of Chi3l1 in either isolated Kupffer cells (KCs, F) or whole liver tissue (Liver, G) from mice fed either NCD or HFHC diet. n=2-3 mice/group. (H) mRNA expression levels of Chil1 in liver tissues of patients with metabolic dysfunction-associated fatty liver (MAFL) or with metabolic dysfunction-associated steatohepatitis (MASH) (GEO Datasets: GSE167523, GSE207310, GSE130970). No-MAFLD or Healthy individuals serve as controls. (I) The correlation between mRNA expression levels of Chil1 and MASLD activity score or fibrosis stage was analyzed (GEO Datasets: GSE130970). Representative images were shown in E. Mann-Whitney test was performed in F. Pearson’s correlation was performed in G. P value and r value are as indicated.

Revised version: (A) Immunofluorescent staining of TIM4 (white), F4/80 (red), Chi3l1 (green), and nuclear DAPI (blue) in liver sections of mice fed with either NCD or HFHC for 16 weeks, illustrating Chi3l1 expression in hepatic macrophages. Scale bar=20 µm and 10 µm (Inset). Chi3l1+ F4/80+ cells/F4/80+ cells were statistically analyzed. n=4 mice/group. (B) Representative immunofluorescence images of liver sections from WT and Chil1-/- mice stained for Chi3l1(green), F4/80 (macrophages), and TIM4 (Kupffer cells). DAPI (blue) marks nuclei. Scale bar=20 µm and 10 µm (Insets). (C, D) Western blot analysis of Chi3l1 in either isolated Kupffer cells (KCs, F) or whole liver tissue (Liver, G) from mice fed either NCD or HFHC diet. n=2-3 mice/group. (E) mRNA expression levels of Chil1 in liver tissues of patients with metabolic dysfunction-associated fatty liver (MAFL) or with metabolic dysfunction-associated steatohepatitis (MASH) (GEO Datasets: GSE167523, GSE207310, GSE130970). No-MAFLD or Healthy individuals serve as controls. (F) The correlation between mRNA expression levels of Chil1 and MASLD activity score or fibrosis stage was analyzed (GEO Datasets: GSE130970). Representative images were shown in A, B. Mann-Whitney test was performed in E. Pearson’s correlation was performed in F. P value and r value are as indicated.

Deficiency of Chi3l1 in Kupffer cells promotes insulin resistance and hepatic lipid accumulation

Earlier version: Given that Chi3l1 is highly expressed in hepatic macrophages, we investigated its functional role by generating mice with conditional knockout (cKO) of Chil1 in either KCs or MoMFs. First, we generated Clec4f △Chil1 mice by crossing Chil1fl/fl mice with Clec4f-cre mice20, achieving KC-specific deletion of Chil1 (Figure S4A-C). These mice, along with Chil1fl/fl controls, were fed either a NCD or a HFHC diet. Under NCD feeding, Clec4fΔChil1 and Chil1fl/fl mice displayed comparable phenotypes in terms of body weight gain, hepatic lipid deposition, metabolic parameters, glucose tolerance and insulin resistance (Figure 2A–F). In contrast, when fed an HFHC diet, Clec4fΔChil1 mice exhibited markedly accelerated weight gain compared to controls (Figure 2A,B). These mice also showed increased hepatic lipid accumulation, as evidenced by H&E and Oil Red O staining at 16 weeks (Figure 2C), along with greater metabolic disturbances, including a higher liver index (liver-to-body weight ratio), elevated serum ALT levels, and increased cholesterol and triglyceride levels in both liver and serum (Figure 2D). Furthermore, Clec4fΔChil1 mice exhibited impaired glucose metabolism, as indicated by worsened glucose tolerance and insulin resistance in IGTT and ITT assays (Figure 2E,F).To exclude potential off-target effects caused by Clec4f-Cre insertion, we compared Clec4f-Cre and Clec4fΔChil1 mice. The phenotypic differences between Clec4f-Cre and Clec4fΔChil1 mice mirrored those observed between Chil1fl/fl and Clec4fΔChil1 mice, with the latter showing faster weight gain, more severe hepatic steatosis, greater metabolic dysregulation, and worsened glucose intolerance and insulin resistance (Figure S5A–G).

Revised version: Given that Chi3l1 is highly expressed in hepatic macrophages, we investigated its functional role by generating mice with conditional knockout (cKO) of Chil1 in either KCs or MoMFs. First, we generated Clec4fΔChil1 mice by crossing Chil1fl/fl mice with Clec4f-cre mice22, achieving KC-specific deletion of Chil1 (Figure S3A-C). These mice, along with Chil1fl/fl controls, were fed either a NCD or a HFHC diet. Under NCD feeding, Clec4fΔChil1 and Chil1fl/fl mice displayed comparable phenotypes in terms of body weight gain, hepatic lipid deposition, metabolic parameters, glucose tolerance and insulin resistance (Figure 2A–F). In contrast, when fed an HFHC diet, Clec4fΔChil1 mice exhibited markedly accelerated weight gain compared to controls (Figure 2A, B). These mice also showed increased hepatic lipid accumulation, as evidenced by H&E and Oil Red O staining at 16 weeks (Figure 2C), along with greater metabolic disturbances, including a higher liver index (liver-to-body weight ratio), elevated serum ALT levels, and increased cholesterol and triglyceride levels in both liver and serum (Figure 2D). Furthermore, Clec4fΔChil1 mice exhibited impaired glucose metabolism, as indicated by worsened glucose tolerance and insulin resistance in IGTT and ITT assays (Figure 2E, F). To exclude potential off-target effects caused by Clec4f-Cre insertion, we compared Clec4f-Cre and Clec4fΔChil1 mice. The phenotypic differences between Clec4f-Cre and Clec4fΔChil1 mice mirrored those observed between Chil1fl/fl and Clec4fΔChil1 mice, with the latter showing faster weight gain, more severe hepatic steatosis, greater metabolic dysregulation, and worsened glucose intolerance and insulin resistance (Figure S4A–G).

Deficiency of Chi3l1 in Kupffer cells promotes insulin resistance and hepatic lipid accumulation.

Earlier version: Chil1fl/fl and Clec4f△Chil1 mice were fed either a normal chow diet (NCD) or a high-fat, high-cholesterol (HFHC) diet for 16 weeks. (A, B) Body weight was recorded during HFHC diet feeding (A) and expressed as a percentage of initial body mass (B). (C) H&E (Upper panel) and oil red o staining (Lower panel) was performed to examine liver histology and hepatic lipid accumulation in in both genotypes after 16 weeks of NCD or HFHC diet. Scale bar = 20 µm. (D) Liver index (liver weight/body weight), ALT levels, and serum and liver Cholesterol or Triglyceride levels were measured in both genotypes after 16 weeks on NCD or HFHC diets. n=4-12 mice/group. (E, F) Intraperitoneal glucose tolerance test (IGTT) and insulin tolerance test (ITT) were performed after 16 weeks of NCD or HFHC feeding in both genotypes (n = 4–12 mice per group). Representative images were shown in (A, E). One-way ANOVA was performed in (A, B, D-F). P-value is as indicated.

Revised version: Chil1fl/fl and Clec4fΔChil1 mice were fed either a normal chow diet (NCD) or a high-fat, high-cholesterol (HFHC) diet for 16 weeks. (A, B) Body weight was recorded during HFHC diet feeding (A) and expressed as a percentage of initial body mass (B). (C) H&E (Upper panel) and oil red o staining (Lower panel) was performed to examine liver histology and hepatic lipid accumulation in both genotypes after 16 weeks of NCD or HFHC diet. Scale bar = 20 µm. (D) Liver index (liver weight/body weight × 100%), ALT levels, and serum and liver Cholesterol or Triglyceride levels were measured in both genotypes after 16 weeks on NCD or HFHC diets. n=4-12 mice/group. (E, F) Intraperitoneal glucose tolerance test (IGTT) and insulin tolerance test (ITT) were performed after 16 weeks of NCD or HFHC feeding in both genotypes (n = 4–12 mice per group). Representative images were shown in (C). One-way ANOVA was performed in (A, B, D-F). P-value is as indicated.

Added in the revised version: To investigate the role of KCs-derived Chi3l1 in MASH, we first examined its expression in a methionine-choline deficient (MCD) diet model. Wildtype mice fed an MCD diet for 6 weeks showed significantly increased Chi3l1 mRNA and protein levels in whole liver tissues compared to NCD controls, confirming diet-induced upregulation (Figure 3A, B). To assess the functional contribution of KCs-derived Chi3l1, we subjected Clec4fΔChil1 mice along with Chil1fl/flcontrols to 6 weeks of MCD diet feeding. Body weight was comparable between genotypes throughout the feeding period (Figure 3C). Histological analysis revealed that loss of Chi3l1 in Kupffer cells led to a significant exacerbation of MCD diet-induced hepatic steatosis, inflammation, and fibrosis, as reflected by increased MASLD activity scores, Oil Red O staining, Sirius Red deposition, and α-SMA expression (Figure 3D). Consistent with this histological finding, Clec4fΔChil1 mice exhibited an increased liver index but similar serum ALT levels, reflecting liver weight gain without evidence of enhanced liver injury (Figure 3E). Additionally, these mice showed significant increases in serum and hepatic triglyceride levels, as well as elevated serum cholesterol, whereas hepatic cholesterol is not significantly upregulated compared to controls (Figure 3E). These data demonstrate that loss of Chi3l1 in Kupffer cells promotes hepatic steatosis, suggesting a protective role for KC-derived Chi3l1 in MASH pathogenesis.

Added in the revised version: Deficiency of Chi3l1 in Kupffer cells promotes liver steatosis and fibrosis in MASH.

Added in the revised version: Male wildtype C57B/6J mice were fed with NCD or MCD diet for 6 weeks (A-B). Chil1fl/fl and Clec4fΔChil1 mice were fed with a MCD diet for 6 weeks (C-E). (A, B) qRT-PCR (A) and wetern blot (B) analysis of Chi3l1 expression in whole liver tissues under NCD and MCD diets. n=3 mice/group. (C) Body weight of mice with conditional deletion of Chil1 in KCs (Clec4fΔChil1) and their control mice (Chil1fl/fl) was recorded during MCD diet. (D) Histological analyses were performed in liver tissue of Clec4fΔChil1 and Chil1fl/flfed the MCD diet for 6 weeks. Scale bar=20μm. (E) Liver index (liver weight/body weight × 100%), ALT levels and serum and liver Cholesterol or Triglyceride levels were measured in both genotypes fed the MCD diet for 6 weeks. n=4-6 mice/group. Representative images are shown in D. Two-tailed, unpaired student t-test was performed in A, C, D, E. P value is as indicated.

Removed in the revised version: Additionally, we constructed another MASH model induced by a methionine and choline deficient diet (MCD), in which extensive inflammatory cell infiltration and fibrosis were induced(Figure S6A, B). Moreover, mRNA and protein levels of Chi3l1 were also upregulated in this model(Figure S6C,D). Although in this model body weight decrease instead of increase, Clec4f △Chil1 mice still showed severe hepatic steatosis and greater metabolic dysregulation (Figure S6E-G). These results further supported that Chi3l1 deficiency in KCs drives MASLD progression.

Earlier version: To assess the role of Chi3l1 in MoMFs, we generated Lyz2△Chil1 mice (Chil1fl/fl × Lyz2-Cre20) and validated Chil1 deletion efficiency in BMDM (Figure S7A–C). Chil1 expression was completely abolished in BMDM from Lyz2ΔChil1 mice. Considering the partial activity of Lyz2-Cre in KCs, we further assessed Chi3l1 expression in KCs isolated from Lyz2Chil1 mice. Only a modest (∼40%) reduction in Chi3l1 mRNA and protein levels was observed in KCs, indicating that Lyz2-Cre–mediated deletion minimally affects Chi3l1 expression in KCs. Lyz2ΔChil1 and Chil1fl/fl control mice were then fed either a NCD or a HFHC diet. Under both dietary conditions, the two genotypes exhibited comparable phenotypes with respect to body weight gain, hepatic lipid accumulation, metabolic parameters, glucose tolerance, and insulin sensitivity (Figure S8A–F). These results indicate that Chi3l1 loss in MoMFs does not substantially impact metabolic regulation.

Revised version: To assess the role of Chi3l1 in MoMFs, we generated Lyz2ΔChil1 mice (Chil1fl/fl × Lyz2-Cre22) and validated Chil1 deletion efficiency in MoMFs and BMDM (Figure S5A–C). Chil1 expression was completely abolished in MoMFs from Lyz2ΔChil1 mice. Considering the partial activity of Lyz2-Cre in KCs, we further assessed Chi3l1 expression in KCs isolated from Lyz2Chil1 mice. Only a modest (∼40%) reduction in Chi3l1 mRNA and protein levels was observed in KCs, indicating that Lyz2-Cre–mediated deletion minimally affects Chi3l1 expression in KCs (Figure S5B–C). Lyz2ΔChil1 and Chil1fl/fl control mice were then fed either a NCD or a HFHC diet. Under both dietary conditions, the two genotypes exhibited comparable phenotypes with respect to body weight gain, hepatic lipid accumulation, metabolic parameters, glucose tolerance, and insulin sensitivity (Figure S6A–F). These results indicate that Chi3l1 loss in MoMFs does not substantially impact metabolic regulation.

ScRNA-seq reveals upregulated glucose metabolism-related transcripts in KCs, correlating with cell death signatures

Added in the revised version: To dissect the distinct metabolic and functional profiles between KCs and MoMFs during MASLD progression, we performed BD Rhapsody single-cell RNA sequencing (scRNA-seq) on liver non-parenchymal cells (NPCs) from mice fed a NCD or HFHC diet for 16 weeks. After quality control and filtration, we retained 23,312 cells from NCD livers and 6,567 cells from HFHC livers for downstream analysis. Using a graph-based clustering approach, we identified 32 distinct cell populations, visualized via uniform manifold approximation and projection (UMAP) (Figure 4A). Monocyte/macrophage subsets were further defined based on lineage-specific markers: Monocytes expressed Ly6c2, Chil3, S100a6, Ccr2, Itgam, and Cx3cr1 but lacked macrophage markers. KCs were marked by Cd68, Vsig4, Clec4f, TIM4, Adgre1, and Clec1b. MoMFs were negative for KCs markers but positive for macrophage markers such as Ccr2, Cx3cr1, Cd9, Itgax, Gpnmb, Cd68, and Adgre1 (Figure 4B, C; UMAP in Figure S7A)8,23.

Removed in the revised version: To dissect the distinct metabolic and functional profiles between KCs and MoMFs during MASLD progression, we analyzed our scRNA-seq data. Consistent with prior studies6,7, we observed decreased KCs numbers but increased MoMFs and monocytes in HFHC-fed mice compared to NCD controls (Figure 3A). Kyoto Encyclopedia of Genes and Genomes (KEGG) pathway analysis showed that while both cell types exhibited activation of phagocytosis-related pathways (lysosome, phagosome, endocytosis, and efferocytosis), they displayed divergent cell fate patterns (Figure 3B). KCs showed strong cell death signatures, whereas MoMFs maintained proliferative activity without evidence of cell death (Figure 3B). Monocytes showed strong cell death and proliferative activity (Figure 3B). Given the significant role of metabolic regulation in cell fate,21 we compared pathways involved in glucose metabolism, cell death, and cell proliferation (Figure 3C). Notably, glucose metabolism pathways were significantly more active in KCs and monocytes compared to MoMFs (Figure 3C). Moreover, the cell proliferation pathway was highly activated in monocytes and consistently activated in MoMFs but not in KCs (Figure 3C). Gene Set Variation Analysis (GSVA)-based correlation analysis revealed a striking association between glucose metabolism and cell death pathways (Figure 3D). These findings demonstrate distinct glucose metabolic activation patterns between KCs and MoMFs, which may underlie their divergent cell fates in MASLD progression.

ScRNA-seq reveals upregulated glucose metabolism-related transcripts in KCs, correlating with cell death signatures.

Earlier version: (A-D) Wildtype C57BL/6J mice were fed either a normal chow diet (NCD) or HFHC for 16 weeks. NPCs were isolated and subjected to BD Rhapsody scRNA sequencing. (A) Quantification of each cell cluster is presented. (B) KEGG analysis reveals the top 12 enriched pathways for up-regulated genes when comparing HFHC versus NCD in KCs, monocytes, and MoMFs, respectively. (C) Gene set variation analysis (GSVA) shows pathway activity for cell death, glucose metabolism, and cell proliferation in KCs, monocytes, and MoMFs of WT mice fed NCD or HFHC for 16 weeks, respectively. (D) The correlation between cell death and glucose metabolism pathways, based on GSVA score, is depicted.

Revised version: Wildtype C57BL/6J mice were fed either a normal chow diet (NCD) or HFHC for 16 weeks. NPCs were isolated and subjected to BD Rhapsody scRNA sequencing. (A) Uniform manifold approximation and projection (UMAP) plots illustrate the clustering of NPCs in the livers of mice fed NCD and HFHC. Cell clusters are color-coded, with monocytes/macrophages clusters outlined. (B) UMAP plots depict the clustering of Monocytes/Macrophages in the livers of mice fed NCD and HFHC. Cell clusters are color-coded. (C) Dot plot displays the scaled gene expression levels of lineage-specific marker genes in different cell clusters. (D) Quantification of each cell cluster is presented. (E) KEGG analysis reveals the top 12 enriched pathways for up-regulated genes when comparing HFHC versus NCD in KCs, monocytes, and MoMFs, respectively. (F) Gene set variation analysis (GSVA) shows pathway activity for cell death, glucose metabolism, and cell proliferation in KCs, monocytes, and MoMFs of WT mice fed NCD or HFHC for 16 weeks, respectively. (G) The correlation between cell death and glucose metabolism pathways, based on GSVA score, is depicted.

Added in the revised version: Consistent with prior studies6,7, we observed decreased KCs numbers but increased MoMFs and monocytes in HFHC-fed mice compared to NCD controls (Figure 4D). Kyoto Encyclopedia of Genes and Genomes (KEGG) pathway analysis showed that while both cell types exhibited activation of phagocytosis-related pathways (lysosome, phagosome, endocytosis, and efferocytosis), they displayed divergent cell fate patterns (Figure 4E). KCs showed strong cell death signatures, whereas MoMFs maintained proliferative activity without evidence of cell death (Figure 4E). Monocytes showed strong cell death and proliferative activity (Figure 4E). Given the significant role of metabolic regulation in cell fate,24 we compared pathways involved in glucose metabolism, cell death, and cell proliferation (Figure 4F). Notably, glucose metabolism pathways were significantly more active in KCs and monocytes compared to MoMFs (Figure 4F). Moreover, the cell proliferation pathway was highly activated in monocytes and consistently activated in MoMFs but not in KCs (Figure 4F). Gene Set Variation Analysis (GSVA)-based correlation analysis revealed a striking association between glucose metabolism and cell death pathways (Figure 4G). These findings demonstrate distinct glucose metabolic activation patterns between KCs and MoMFs, which may underlie their divergent cell fates in MASLD progression.

Chi3l1 deficiency promote KCs death during MASLD

Earlier version: To investigate the role of Chi3l1 in KCs survival during MASLD, we generated Chil1-/- mice, as Chi3l1 is a secreted protein22,23. The successful knockout was confirmed by qRT-PCR analysis of liver tissue (Figure S9A, B). We then performed scRNA-seq on NPCs isolated from Chil1-/- mice fed an HFHC diet for 16 weeks. After quality control, 6,813 high-quality cells were retained for analysis. Using established KC markers (Figure 1C), we conducted GSVA to examine metabolic pathways. This revealed enhanced cell death pathways in KCs from HFHC-fed mice, with significantly greater apoptosis signatures in Chil1-/- KCs compared to wild-type (WT) controls (Figure 4A). The increased apoptosis was further supported by upregulation of pro-apoptotic genes in Chil1-/- KCs (Figure 4B).

Revised version: To investigate the role of Chi3l1 in KCs survival during MASLD, we then performed scRNA-seq on NPCs isolated from Chil1-/- mice fed an HFHC diet for 16 weeks. After quality control, 6,813 high-quality cells were retained for analysis. Using established KC markers (Figure 4C), we conducted GSVA to examine metabolic pathways. This revealed enhanced cell death pathways in KCs from HFHC-fed mice, with significantly greater apoptosis signatures in Chil1-/- KCs compared to wild-type (WT) controls (Figure 5A). The increased apoptosis was further supported by upregulation of pro-apoptotic genes in Chil1-/-KCs (Figure 5B).

Chi3l1 deficiency promote KCs death during MASLD.

Earlier version: (A) GSVA analysis showed the enrichment of cell death-related pathways in KCs from WT mice fed with either NCD or HFHC or Chil1-/- mice fed with HFHC. (B) Dot plot showing the scaled gene expression levels of Apoptosis-related genes and repressor genes in KCs from either WT or Chil1-/- fed with HFHC. (C) Strategy used to gate KCs (CD45+ F4/80hi CD11blow TIM4hi) and MoMFs (CD45+ F4/80low CD11bhi Ly6G− TIM4−) in the liver by flow cytometry. (D) Number of KCs and MoMFs /liver or gram(g) liver were statistically analyzed. n= 3-4 mice per group. (E) Immunofluorescent staining to detect TIM4(green), TUNEL (red), and nuclear DAPI (blue) in liver sections. Scale bar=20μm and 5μm(zoom). TUNEL+ TIM4+ cells/TIM4+ cells were statistically analyzed. n=4-6 mice/group. Representative images are shown in C, E. One-way ANOVA was performed in D. Two-tailed, unpaired student t-test was performed in E. P value is as indicated.

Revised version: (A) GSVA analysis showed the enrichment of cell death-related pathways in KCs from WT mice fed with either NCD or HFHC or Chil1-/- mice fed with HFHC. (B) Dot plot showing the scaled gene expression levels of Apoptosis-related genes and repressor genes in KCs from either WT or Chil1-/-fed with HFHC. (C) Strategy used to gate KCs (CD45+ F4/80hi CD11blow TIM4hi) and MoMFs (CD45+ F4/80low CD11bhi Ly6G- TIM4-) in the liver by flow cytometry. (D) Number of KCs and MoMFs /liver or gram(g) liver were statistically analyzed. n= 3-4 mice per group. (E) Immunofluorescent staining to detect TIM4(green), TUNEL (red), and nuclear DAPI (blue) in liver sections. Scale bar=50µm and 20µm (Insets). TUNEL+ TIM4+ cells/TIM4+ cells were statistically analyzed. n=4 mice/group. Representative images are shown in C, E. One-way ANOVA was performed in D. Two-tailed, unpaired student t-test was performed in E. P value is as indicated.

Earlier version: We next validated these findings by flow cytometry using the gating strategy shown in Figure 4C. While WT and Chil1-/- mice showed similar KC numbers at baseline, dramatic differences emerged during HFHC feeding. WT KC numbers remained stable at 8 weeks but decreased by 50% at 16 weeks. In contrast, Chil1-/- mice exhibited accelerated KCs loss, with a 30% reduction by 8 weeks progressing to 60% by 16 weeks (Figure 4D, S10A). Notably, MoMFs populations remained comparable between groups at early timepoints but showed greater reduction in Chil1-/- mice at 16 weeks ((Figure 4D, S10A).

Revised version: We next validated these findings by flow cytometry using the gating strategy shown in Figure 5C. While WT and Chil1-/- mice showed similar KC numbers at baseline, dramatic differences emerged during HFHC feeding. WT KC numbers remained stable at 8 weeks but decreased by 50% at 16 weeks. In contrast, Chil1-/- mice exhibited accelerated KCs loss, with a 30% reduction by 8 weeks progressing to 60% by 16 weeks (Figure 5D, S8A). Notably, MoMFs populations remained comparable between groups at early timepoints but showed greater reduction in Chil1-/- mice at 16 weeks (Figure 5D, S8A).

Earlier version: Histological analysis further supported these findings. TIM4/TUNEL co-staining revealed no TUNEL+ KCs in WT livers at baseline, whereas 40% and 50% of KCs were TUNEL+ at 8 and 16 weeks, respectively. In Chil1-/- mice, KC apoptosis was significantly increased at both time points (Figure 4E). Consistent results were obtained with TIM4/cleaved caspase-3 co-staining (Figure S10B). We further confirmed these observations in Clec4fΔChil1 mice in both HFHC24 and MCD diet models. In the MCD model, Clec4fΔChil1 mice exhibited enhanced KCs death compared with Chil1fl/fl controls (Figure S11A). To exclude potential effects of myeloid cell–derived Chi3l1 on KCs survival, we compared KCs death and abundance between Chil1fl/fl and Lyz2ΔChil1 mice using histological and flow cytometric analyses. Loss of Chi3l1 in MoMFs did not lead to significant KC apoptosis or depletion (Figure S11B–D).Together, these results demonstrate that Chi3l1 deficiency promotes KC apoptosis, resulting in premature KC depletion during MASLD progression.

Revised version: Histological analysis further supported these findings. TIM4/TUNEL co-staining revealed no TUNEL⁺ KCs in WT livers at baseline, whereas 40% and 50% of KCs were TUNEL⁺ at 8 and 16 weeks, respectively. In Chil1-/- mice, KC apoptosis was significantly increased at both time points (Figure 5E). Consistent results were obtained with TIM4/cleaved caspase-3 co-staining (Figure S8B). We further confirmed these observations in Clec4fΔChil1 mice in both HFHC25 and MCD diet models. In the MCD model, Clec4fΔChil1 mice exhibited enhanced KCs death compared with Chil1fl/fl controls (Figure S9A). To exclude potential effects of myeloid cell–derived Chi3l1 on KCs survival, we compared KCs death and abundance between Chil1fl/fl and Lyz2ΔChil1 mice using histological and flow cytometric analyses. Loss of Chi3l1 in MoMFs did not lead to significant KC apoptosis or depletion (Figure S9B–D). Together, these results demonstrate that Chi3l1 deficiency promotes KC apoptosis, resulting in premature KC depletion during MASLD progression.

Molecular interaction between Chi3l1 and glucose

Earlier version: Our investigation into Chi3l1-mediated KCs survival revealed an unexpected structural relationship: Chi3l1 binds to glucose, which is structurally analogous to chitin, a polysaccharide well known to bind Chi3l1(Figure 5A). Bioinformatics analysis using the STITCH database further supported this observation, predicting a high probability of direct Chi3l1-glucose interaction (Figure 5B). To experimentally validate this interaction, we performed pull-down assays using biotin-labeled glucose incubated with plasma from HFHC-fed mice. Streptavidin bead isolation followed by anti-Chi3l1 Western blotting demonstrated specific binding between Chi3l1 and biotin-glucose, but not biotin alone (Figure 5C, D). This interaction was competitively inhibited by unlabeled glucose, confirming specificity (Figure 5D). Quantitative analysis using microscale thermophoresis with recombinant mouse Chi3l1 (rChi3l1) yielded a dissociation constant (Kd) of 4.95 mM for the Chi3l1-glucose interaction (Figure 5E). Notably, circulating Chi3l1 levels were significantly elevated in serum from HFHC-fed mice compared to baseline (Figure 5F), suggesting a potential physiological role for this interaction in metabolic regulation. These findings establish Chi3l1 as a novel glucose-binding protein that may participate in glucose homeostasis during MASLD progression.

Revised version: Our investigation into Chi3l1-mediated KCs survival revealed an unexpected structural relationship: Chi3l1 binds to glucose, which is structurally analogous to chitin, a polysaccharide well known to bind Chi3l1 (Figure 6A). Bioinformatics analysis using the STITCH database further supported this observation, predicting a high probability of direct Chi3l1-glucose interaction (Figure 6B). To experimentally validate this interaction, we performed pull-down assays using biotin-labeled glucose incubated with plasma from HFHC-fed mice. Streptavidin bead isolation followed by anti-Chi3l1 Western blotting demonstrated specific binding between Chi3l1 and biotin-glucose, but not biotin alone (Figure 6C, D). This interaction was competitively inhibited by unlabeled glucose, confirming specificity (Figure 6D). Quantitative analysis using microscale thermophoresis with recombinant mouse Chi3l1 (rChi3l1) yielded a dissociation constant (Kd) of 4.95 mM for the Chi3l1-glucose interaction (Figure 6E). Notably, circulating Chi3l1 levels were significantly elevated in serum from HFHC-fed mice compared to baseline (Figure 6F), suggesting a potential physiological role for this interaction in metabolic regulation. These findings establish Chi3l1 as a novel glucose-binding protein that may participate in glucose homeostasis during MASLD progression.

Molecular interaction between Chi3l1 and glucose.

(A) A comparison of chemical structures between glucose and chitin. (B) Prediction of Chi3l1-glucose interaction using STITCH database (http://stitch.embl.de). (C) Strategy for pulling down glucose-binding proteins in murine serum. (D) Biotin-conjugated glucose was incubated with murine serum from mice fed with HFHC for 16 weeks. Proteins bound to glucose were precipitated by streptavidin beads. Biotin or biotin-conjugated glucose plus glucose were used as negative controls. Western blot was performed to examine Chi3l1 in the precipitate. (E) Microscale thermophoresis assay to detect the interaction between recombinant mouse Chi3l1 (rChi3l1) and glucose. Kd=4.95±0.66mM. (F) Western blot to detect Chi3l1 expression in murine serum before and after HFHC feeding. n=3 mice/group.

Chi3l1 limits glucose uptake and protects hepatic macrophages from cell death

Earlier version: To elucidate the functional consequences of Chi3l1-glucose binding, we examined glucose metabolism in hepatic macrophages. Using the fluorescent glucose analog 2-NBDG25, we performed uptake assays in KCs following 12-hour glucose starvation. While glycogen droplet size remained unchanged in untreated KCs regardless of rChi3l1 supplementation (Figure 6A), 2-NBDG exposure significantly increased glycogen accumulation. This effect was markedly suppressed by rChi3l1 co-treatment (Figure 6A), a phenotype replicated in BMDM (Figure 6A). These results demonstrate that Chi3l1 restricts glucose uptake and subsequent glycogen storage.

Revised version: To elucidate the functional consequences of Chi3l1-glucose binding, we examined glucose metabolism in hepatic macrophages. Using the fluorescent glucose analog 2-NBDG26, we performed uptake assays in KCs following 12-hour glucose starvation. While glycogen droplet size remained unchanged in untreated KCs regardless of rChi3l1 supplementation (Figure 7A), 2-NBDG exposure significantly increased glycogen accumulation. This effect was markedly suppressed by rChi3l1 co-treatment (Figure 7A), a phenotype replicated in BMDM (Figure 7A). These results demonstrate that Chi3l1 restricts glucose uptake and subsequent glycogen storage.

Chi3l1 limits glucose uptake and protects hepatic macrophages from cell death.

Earlier version: (A) Following 12 h of glucose starvation, isolated KCs or BMDM were divided into two groups: one treated with no 2-NBDG and the other with 2-NBDG. Within each group, KCs or BMDM were further treated without or with recombinant murine Chi3l1 (rChi3l1) for 6 h. Glycogen aggregate formation labeled by 2-NBDG (Green) in KCs or BMDM was examined after counterstaining with nuclear DAPI (Blue). Scale bar=2μm. Area of 2-NBDG in KCs were quantified. (B) Following 12 h of glucose starvation, BMDM were treated with either no glucose or high glucose (25mM). Concurrently, BMDM were treated without or with rChi3l1 for 24 h under each condition. glycogen aggregate formation in BMDM was detected using immunofluorescence staining for Stbd1 (red) and nuclear DAPI (blue). Scale bar = 10 μm. (C and D) BMDM cells were treated without or with rChi3l1 for 24 h and subjected to Seahorse metabolic analysis to measure the extracellular acidification rate (ECAR). (E and F) KCs were treated without (blank) or with either Isopropyl alcohol(Iso) or 800uM palmitic acid (PA) or 100ng rChi3l1 with 800 uM PA for 24 h. Western blot was performed to detect cleaved caspase 3 (Cl-Casp3) in E. Calcein/PI staining was quantified to detect cell viability in F. Scale bar=50μm. (G) Measurement of 2-NBDG (a fluorescent glucose analog) uptake by KCs in vivo. WT and Chil1-/- mice, either untreated or supplemented with rChi3l1, were injected intraperitoneally with 12 mg/kg 2-NBDG. After 45mins, KCs were isolated and glucose uptake assessed by spectrophotometry. (H) Representative immunofluorescence images of liver sections stained for TIM4 (red) and 2-NBDG uptake (green) to visualize glucose uptake by KCs in situ. Scale bar = 10 µm (zoom). Quantification is shown as the percentage of TIM4+ cells that are also 2-NBDG+. Representative images were shown in A, B, H. One-way ANOVA was performed in A, F, G, H. Two-tailed, unpaired student t-test was performed in D. P value is as indicated.

Revised version: (A) Following 12 h of glucose starvation, isolated KCs or BMDM were divided into two groups: one treated with no 2-NBDG and the other with 2-NBDG. Within each group, KCs or BMDM were further treated without or with recombinant murine Chi3l1 (rChi3l1) for 6 h. Glycogen aggregate formation labeled by 2-NBDG (Green) in KCs or BMDM was examined after counterstaining with nuclear DAPI (Blue). Scale bar=2μm. Area of 2-NBDG in KCs were quantified. (B) Following 12 h of glucose starvation, BMDM were treated with either no glucose or high glucose (25mM). Concurrently, BMDM were treated without or with rChi3l1 for 24 h under each condition. glycogen aggregate formation in BMDM was detected using immunofluorescence staining for Stbd1 (red) and nuclear DAPI (blue). Scale bar = 10 μm. (C and D) BMDM cells were treated without or with rChi3l1 for 24 h and subjected to Seahorse metabolic analysis to measure the extracellular acidification rate (ECAR). (E and F) KCs were treated without (blank) or with either Isopropyl alcohol (Iso) or 800uM palmitic acid (PA) or 100ng rChi3l1 with 800 uM PA for 24 h. Western blot was performed to detect cleaved caspase 3 (Cl-Casp3) in E. Calcein/PI staining was quantified to detect cell viability in F. Scale bar=50μm. (G) Measurement of 2-NBDG (a fluorescent glucose analog) uptake by KCs in vivo. WT and Chil1-/-mice, either untreated or supplemented with rChi3l1, were injected intraperitoneally with 12 mg/kg 2-NBDG. After 45mins, KCs were isolated and glucose uptake assessed by spectrophotometry. (H) Representative immunofluorescence images of liver sections stained for TIM4 (red) and 2-NBDG uptake (green) to visualize glucose uptake by KCs in situ. Scale bar = 10 µm (Insets). Quantification is shown as the percentage of TIM4⁺ cells that are also 2-NBDG⁺. Representative images were shown in A, B, H. One-way ANOVA was performed in A, F, G, H. Two-tailed, unpaired student t-test was performed in D. P value is as indicated.

Earlier version: Further validation using Stbd1 (a glycogen-binding protein25) immunofluorescence revealed minimal glycogen foci in glucose-deprived BMDM, with no rChi3l1-dependent differences. High-glucose conditions, however, triggered robust glycogen aggregation, which was significantly attenuated by rChi3l1 (Figure 6B). Concordantly, extracellular acidification rate (ECAR) measurements showed reduced basal and total glycolytic capacity in rChi3l1-treated BMDMs (Figure 6C, D), confirming Chi3l1’s role in limiting glucose metabolism.

Revised version: Further validation using Stbd1 (a glycogen-binding protein26) immunofluorescence revealed minimal glycogen foci in glucose-deprived BMDM, with no rChi3l1-dependent differences. High-glucose conditions, however, triggered robust glycogen aggregation, which was significantly attenuated by rChi3l1 (Figure 7B). Concordantly, extracellular acidification rate (ECAR) measurements showed reduced basal and total glycolytic capacity in rChi3l1-treated BMDMs (Figure 7C, D), confirming Chi3l1’s role in limiting glucose metabolism.

Earlier version: To test whether Chi3l1-glucose binding influence cell survival, we employed a palmitic acid (PA)-induced lipotoxicity cell-based model to better mimic the in vivo environment. rChi3l1 supplementation reduced PA-induced cleavage of caspase-3 (Figure 6E) and decreased KCs death (calcein/PI staining, Figure 6F). To validate this mechanism in vivo, we intraperitoneally injected 2-NBDG into WT and Chil1-/- mice, with or without supplementation of rChi3l1, to assess glucose uptake by KCs. Chil1-/- KCs displayed markedly increased 2-NBDG uptake compared with WT controls, whereas rChi3l1 supplementation significantly reduced glucose uptake. These results demonstrate that serum Chi3l1 limits glucose uptake by KCs in vivo (Figure 6G, H).Collectively, these findings demonstrate that Chi3l1 protects KCs from metabolic stress–induced death by regulating glucose uptake.

Revised version: To test whether Chi3l1-glucose binding influence cell survival, we employed a palmitic acid (PA)-induced lipotoxicity cell-based model to better mimic the in vivo environment. rChi3l1 supplementation reduced PA-induced cleavage of caspase-3 (Figure 7E) and decreased KCs death (calcein/PI staining, Figure 7F). To validate this mechanism in vivo, we intraperitoneally injected 2-NBDG into WT and Chil1-/- mice, with or without supplementation of rChi3l1, to assess glucose uptake by KCs. Chil1-/- KCs displayed markedly increased 2-NBDG uptake compared with WT controls, whereas rChi3l1 supplementation significantly reduced glucose uptake. These results demonstrate that serum Chi3l1 limits glucose uptake by KCs in vivo (Figure 7G, H). Collectively, these findings demonstrate that Chi3l1 protects KCs from metabolic stress–induced death by regulating glucose uptake.

Discussion

Earlier version: Our findings establish Chi3l1 as a critical metabolic regulator that controls hepatic macrophage fate through a novel glucose-dependent mechanism in MASLD. Using cell-specific knockout models, we uncovered a fundamental dichotomy in Chi3l1 function: selective ablation in KCs dramatically accelerated MASLD progression and metabolic dysfunction, whereas deletion in MoMFs produced minimal metabolic effects. Single-cell transcriptomics revealed the molecular basis for this cell-type specificity - KCs exhibit a glucose-hungry metabolic phenotype that renders them uniquely dependent on Chi3l1-mediated regulation, while MoMFs maintain a relatively glucose-independent metabolic program. At the mechanistic level, we demonstrate that Chi3l1 functions as a physiological glucose sensor, directly binding extracellular glucose to limit its cellular uptake. This interaction establishes a crucial metabolic safeguard that specifically protects glucose-dependent KCs from lethal metabolic stress while sparing glucose-independent MoMFs. Through this precise modulation of glucose availability, Chi3l1 maintains metabolic homeostasis and preserves KCs populations during chronic dietary challenge (Figure 7).

Revised version: Our findings establish Chi3l1 as a critical metabolic regulator that controls hepatic macrophage fate through a novel glucose-dependent mechanism in MASLD. Using cell-specific knockout models, we uncovered a fundamental dichotomy in Chi3l1 function: selective ablation in KCs dramatically accelerated MASLD progression and metabolic dysfunction, whereas deletion in MoMFs produced minimal metabolic effects. Single-cell transcriptomics revealed the molecular basis for this cell-type specificity - KCs exhibit a glucose-hungry metabolic phenotype that renders them uniquely dependent on Chi3l1-mediated regulation, while MoMFs maintain a relatively glucose-independent metabolic program. At the mechanistic level, we demonstrate that Chi3l1 functions as a physiological glucose sensor, directly binding extracellular glucose to limit its cellular uptake. This interaction establishes a crucial metabolic safeguard that specifically protects glucose-dependent KCs from lethal metabolic stress while sparing glucose-independent MoMFs. Through this precise modulation of glucose availability, Chi3l1 maintains metabolic homeostasis and preserves KCs populations during chronic dietary challenge (Figure 8).

Differential regulation of KCs and MoMFs fate by Chi3l1-glucose interaction.

KCs maintain a high-glucose activation state, while MoMFs exhibit a relatively low-glucose metabolic program. Chi3l1-glucose binding inhibits glucose uptake in KCs, thereby delaying KCs death and alleviating MASLD progression and metabolic dysfunction. In contrast, although Chi3l1-glucose binding similarly inhibits glucose uptake in MoMFs, their low basal glucose metabolism renders them resistant to this metabolic perturbation, resulting in minimal impact on MASLD pathogenesis.

Analysis of publicly available scRNA-seq datasets, including those from the Liver Atlas and prior studies7–9, indicates that Chil1 transcripts are mainly detected in neutrophils. In contrast, our immunofluorescence data show that Chi3l1 protein is predominantly localized in Kupffer cells under normal conditions and in both KCs and MoMFs during MASLD progression. This discrepancy likely reflects differences in transcript versus protein abundance and detection sensitivity. While scRNA-seq captures relative mRNA levels per cell, tissue-based staining reflects both expression and cell prevalence, highlighting macrophages as a major contributor to total hepatic Chi3l1 protein. Moreover, environmental factors such as diet, microbiota, or disease stage may influence Chil1 expression patterns across immune cell types.

Earlier version: Our study reveals fundamental differences in metabolic requirements between hepatic macrophage subsets that provide new insights into MASLD pathogenesis. We demonstrate that KCs and MoMFs play stage-specific roles in disease progression, with KCs serving as critical regulators of early metabolic homeostasis while MoMFs appear more involved in later inflammatory phases. This temporal specialization explains the striking dichotomy observed in our genetic models - KCs-specific Chi3l1 deletion dramatically exacerbated metabolic dysfunction, whereas MoMFs deletion showed minimal effects. The heightened glucose metabolism of KCs during MASLD renders them uniquely vulnerable to dietary stress. Chi3l1 serves as a crucial metabolic buffer in this context, directly protecting KCs through glucose modulation as evidenced by reduced glycogen accumulation and attenuated glycolytic flux. Our findings using the HFHC model complement previous findings in fibrogenic CDAA-HFAT models26 or MCD/CCL4 models27 or human livers28, collectively suggesting Chi3l1 may have dual roles in MASLD - maintaining metabolic balance through KCs in early disease while potentially influencing fibrogenesis via MoMFs in advanced stages. The accelerated KCs death in knockout models provides direct experimental evidence linking macrophage survival to metabolic outcomes, resolving key questions about MASLD progression mechanisms.

Revised version: Our study reveals fundamental differences in metabolic requirements between hepatic macrophage subsets that provide new insights into MASLD pathogenesis. We demonstrate that KCs and MoMFs play stage-specific roles in disease progression, with KCs serving as critical regulators of early metabolic homeostasis while MoMFs appear more involved in later inflammatory phases. This temporal specialization explains the striking dichotomy observed in our genetic models-KCs-specific Chi3l1 deletion dramatically exacerbated metabolic dysfunction, whereas MoMFs deletion showed minimal effects. The heightened glucose metabolism of KCs during MASLD renders them uniquely vulnerable to dietary stress. Chi3l1 serves as a crucial metabolic buffer in this context, directly protecting KCs through glucose modulation as evidenced by reduced glycogen accumulation and attenuated glycolytic flux. Our findings using the HFHC model complement previous findings in fibrogenic CDAA-HFAT models27 or MCD/CCL4 models28 or human livers29, collectively suggesting Chi3l1 may have dual roles in MASLD - maintaining metabolic balance through KCs in early disease while potentially influencing fibrogenesis via MoMFs in advanced stages. The accelerated KCs death in knockout models provides direct experimental evidence linking macrophage survival to metabolic outcomes, resolving key questions about MASLD progression mechanisms.

Earlier version: The structural characteristics of Chi3l1 have been extensively studied. Chi3l1 forms a homodimer, with each subunit containing a catalytic domain and a carbohydrate-binding domain. While the catalytic domain retains structural similarity to chitinases, it lacks enzymatic activity,29,30 and the carbohydrate-binding domain mediates interactions with carbohydrate ligands.29 While chitin-binding domains are traditionally known to interact with complex polysaccharides, our findings reveal that Chi3l1 (YKL-40), a mammalian chitinase-like protein, specifically binds to glucose—a simple monosaccharide. This represents a fundamental departure from canonical binding to insoluble polymers such as chitin and suggests a previously unrecognized role for Chi3l1 in monosaccharide recognition, potentially linking it to glucose metabolism and energy sensing. Furthermore, we observed that Chi3l1 protein levels increased in the serum of mice fed a high-fat, high-cholesterol (HFHC) diet for 16 weeks (Figure 5F) but plateaued with prolonged feeding (24 weeks; data not shown), suggesting an adaptive regulatory limit. Together, these findings indicate that Chi3l1 possesses glucose-binding capacity that may be functionally relevant but limited in vivo.

Revised version: The structural characteristics of Chi3l1 have been extensively studied. Chi3l1 forms a homodimer, with each subunit containing a catalytic domain and a carbohydrate-binding domain. While the catalytic domain retains structural similarity to chitinases, it lacks enzymatic activity,30,31 and the carbohydrate-binding domain mediates interactions with carbohydrate ligands.30 While chitin-binding domains are traditionally known to interact with complex polysaccharides, our findings reveal that Chi3l1 (YKL-40), a mammalian chitinase-like protein, specifically binds to glucose—a simple monosaccharide. This represents a fundamental departure from canonical binding to insoluble polymers such as chitin and suggests a previously unrecognized role for Chi3l1 in monosaccharide recognition, potentially linking it to glucose metabolism and energy sensing. Furthermore, we observed that Chi3l1 protein levels increased in the serum of mice fed a high-fat, high-cholesterol (HFHC) diet for 16 weeks (Figure 6F) but plateaued with prolonged feeding (24 weeks; data not shown), suggesting an adaptive regulatory limit. Together, these findings indicate that Chi3l1 possesses glucose-binding capacity that may be functionally relevant but limited in vivo.

Earlier version: Our findings carry important translational potential for MASLD treatment. The discovery of Chi3l1’s glucose-sensing function in KCs suggests two complementary therapeutic strategies: first, developing Chi3l1-based interventions to preserve KC populations during early metabolic dysfunction; second, creating cell-type-specific approaches that selectively modulate glucose metabolism in KCs while sparing MoMFs. Importantly, although access to early-stage human liver tissue is limited due to the asymptomatic nature of the disease, multiple human studies have consistently reported elevated Chi3l1 levels in steatotic and fibrotic liver disease27,31,32, underscoring the clinical relevance of our mechanistic findings. Building on this evidence, the structural mapping of Chi3l1’s glucose-binding domain now enables rational design of small-molecule mimetics or biologics to therapeutically enhance this protective pathway. Besides, several key questions emerge for future research to advance these therapeutic possibilities: (1) How glucose levels are coordinated with other death inducers such as lipid toxicity; (2) Whether competing carbohydrate ligands modulate Chi3l1’s glucose-sensing capacity in different metabolic states; (3) The clinical relevance of human Chi3l1 variants in MASLD susceptibility and progression. Addressing these questions will be crucial for translating our mechanistic insights into targeted therapies that account for the complex metabolic specialization of hepatic macrophage subsets.

Revised version: Our findings carry important translational potential for MASLD treatment. The discovery of Chi3l1’s glucose-sensing function in KCs suggests two complementary therapeutic strategies: first, developing Chi3l1-based interventions to preserve KC populations during early metabolic dysfunction; second, creating cell-type-specific approaches that selectively modulate glucose metabolism in KCs while sparing MoMFs. Importantly, although access to early-stage human liver tissue is limited due to the asymptomatic nature of the disease, multiple human studies have consistently reported elevated Chi3l1 levels in steatotic and fibrotic liver disease28,32,33, underscoring the clinical relevance of our mechanistic findings. Building on this evidence, the structural mapping of Chi3l1’s glucose-binding domain now enables rational design of small-molecule mimetics or biologics to therapeutically enhance this protective pathway. Besides, several key questions emerge for future research to advance these therapeutic possibilities: (1) How glucose levels are coordinated with other death inducers such as lipid toxicity; (2) Whether competing carbohydrate ligands modulate Chi3l1’s glucose-sensing capacity in different metabolic states; (3) Functional validation in primary human macrophages or human liver tissues would further strengthen the translational significance of this work. Addressing these questions will be crucial for translating our mechanistic insights into targeted therapies that account for the complex metabolic specialization of hepatic macrophage subsets.

Our findings reveal a novel metabolic checkpoint in which Chi3l1 selectively sustains KCs populations by modulating glucose metabolism, offering key insights into MASLD pathogenesis. The study highlights the therapeutic potential of targeting Chi3l1-glucose interactions to preserve protective KCs and curb MASLD progression. Future research should explore whether Chi3l1 supplementation or pharmacological modulation can rescue KCs viability, as well as investigate whether this mechanism extends to other macrophage-driven metabolic disorders, such as MASH or diabetes. By identifying cell type-specific metabolic vulnerabilities, this work paves the way for precision therapies that selectively manipulate macrophage subsets to treat liver disease.